# Homework 0: Introduction & Numpy Basics
**Submission Deadline: May 08, 2026, Friday, 9:30**

---

> **Machine Learning - Foundations and Algorithms** in Summer Semester 2026
>
> - Prof. Gerhard Neumann
> - Instructors for this exercise: Florian Seligmann (florian.seligmann@kit.edu) and Serge Thilges (serge.thilges@kit.edu)
> 
> ⚠️ For general questions about the exercises, please post in the Ilias forum so we can answer them once for all students. Only use email for individual questions whose answers are not relevant to all students.

---

## Homework Info

For this exercise, please try to fill in the code where stated.
It demonstrates some basic operations we will need for this lecture.
**Then, submit this exercise following the instructions above.**

**This submission is only for getting familiar with the submission procedure.
It will not be graded.**

## Group Submission

- The exercise sheets can be submitted in **groups of up to 3 students**.
- **Each person in the group must upload the edited notebook to Ilias**. It is not sufficient if only one person from the group does this. If you forgot to upload but at least on other group member upholds his solution, you will receive 2/3 of the points the first time this happens, and zero points if it happens again.
- **The submission deadline is strict**. Late submissions receive zero points. **We do not accept submissions via e-mail**. 
- It is possible to join a new group during the semester if your own group dissolves early. Each group must upload their own solution, and we will check submissions for duplicates.
- You can use the team building forum (https://ilias.studium.kit.edu/ilias.php?baseClass=ilrepositorygui&cmdNode=xs:lz:5y&cmdClass=ilBookingObjectGUI&ref_id=2930523) to team up. If you already have group members or want to submit alone, you do not need to use the forum.


Groups are automatically recorded for each exercise, so **please enter the u-identifiers of your group members in the *next* cell.** If your group consists of only 2 students or if you are submitting alone, leave the remaining fields empty.

Here is an **example** for a group consisting of uabcd and uefgh:

_U-identifiers of group members:_

_Member 1: uabcd_

_Member 2: uefgh_

_Member 3:_

Enter your group members in the *next* cell!

U-identifiers of group members:

Member 1:

Member 2:

Member 3:

## Auto-grading

We use an auto-grading system that automatically analyzes your submitted Jupyter Notebooks and checks them for correctness using hidden tests. These tests determine the points you receive for the exercise sheet.

To ensure auto-grading works smoothly, please note the following:

- The notebook must have the filename "ex_00_intro.ipynb"
- **Upload only the Jupyter Notebook** on Ilias as a single file (not as a zip!). **Do not include any additional files**.
- Before submitting a notebook, test that everything runs without errors from start to finish.
- **Cells marked with "##### DO NOT CHANGE #####" must not be edited or deleted**
  - There may be seemingly empty cells that are also marked with "##### DO NOT CHANGE #####". These must also not be edited or deleted.
  - If you do modify them, auto-grading will not work and you will receive no points.
  - We will be strict about this and make no exceptions if someone modifies cells clearly marked as read-only!
- Your solution must be entered in the correct cell (marked with "# YOUR CODE HERE").
  - Please **delete the NotImplementedError**
- Generally, **do not delete any cells** and **do not add any cells**. The cells where your solution should be entered already exist (marked with "# YOUR CODE HERE").
- The Jupyter notebooks have inline tests (visible to you) that check your result for basic correctness.
  - These are primarily for you to identify and correct errors.
  - However, the inline tests you can see in the notebook are not the tests used for grading!
  - Passing the inline tests is a necessary but not sufficient condition to receive points!

**If you fail to comply with these rules, your submission may not be graded and you may receive no points for this exercise!**

## Exercise 1: Matrices and Vectors in Python
As mentioned, we will use Python for the coding exercise. We will use version 3.8.
Additionally, for this and all following exercise we will use NumPy, one of the most fundamental python libraries,
designed to efficiently create and operate on multi-dimensional arrays (, i.e., vectors, matrices and higher order tensors)

If you are new to NumPy we refer to https://numpy.org/devdocs/user/quickstart.html .


### Creatin Matrices and Vectors with numpy
We first create a matrix <br>
\begin{align*}A =
\begin{pmatrix}
    1 & 2 & 3 \\
    4 & 2 & 6 \\
    7 & 3 & 8
\end{pmatrix}
\end{align*}
and two vectors <br>

\begin{align*} v =
\begin{pmatrix}
    3  \\
    1  \\
    2
\end{pmatrix}, ~w =
\begin{pmatrix}
    1  \\
    1  \\
    2
\end{pmatrix}
\end{align*}



```python
##### DO NOT CHANGE #####
import numpy as np

# Let us create a matrix A and vectors v and w:
A = np.array([[1, 2, 3],
              [4, 2, 6],
              [7, 3, 8]])
v = np.array([3, 1, 2])
w = np.array([1, 1, 2])
print("Matrix A:")
print(A)
print("Vector v:")
print(v)
print("Vector w:")
print(w)



##### DO NOT CHANGE #####
```

### Task 1.1: Element-Wise Operations

Implement the element-wise sum v + w and element-wise multiplication 2 $\cdot$ A. Print out your results and verify that they are correct.

<!-- grade_id: task_1_1, points: 0, grade: false, solution: true -->
```python
elem_wise_sum = None
elem_wise_mult = None

# Do the element-wise sum v+w here
# YOUR CODE HERE
raise NotImplementedError()

print(elem_wise_sum)
print("----------")

# Do the element-wise multiblication 2*A here
# YOUR CODE HERE
raise NotImplementedError()

print(elem_wise_mult)
```

<!-- grade_id: test_task_1_1, points: 0, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_1_1 - possible points: 0

# Test that all shapes are correct
assert elem_wise_sum.shape == (3,)
assert elem_wise_mult.shape == (3, 3)

print("✓ Shapes are correct")


##### DO NOT CHANGE #####
```

### Task 1.2: Matrix-Vector and Matrix-Matrix Products

Do the matrix vector product $A \cdot v$ and the matrix matrix product $A \cdot A$. Print out your result and verify that the result is correct. To perform a matrix multiplication, use either `np.dot` or just the `@` symbol

<!-- grade_id: task_1_2, points: 0, grade: false, solution: true -->
```python
# 

matrix_vector_prod = None
matrix_matrix_prod = None

# Do the matrix vector product A * v here
# YOUR CODE HERE
raise NotImplementedError()

print(matrix_vector_prod)
print("----------")

# Do the matrix matrix product A * A here
# YOUR CODE HERE
raise NotImplementedError()

print(matrix_matrix_prod)
```

<!-- grade_id: test_task_1_2, points: 0, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_1_2 - possible points: 0

# Test that all shapes are correct
assert matrix_vector_prod.shape == (3,)
assert matrix_matrix_prod.shape == (3, 3)

print("✓ Shapes are correct")


##### DO NOT CHANGE #####
```

## Exercise 2: Broadcasting 

Broadcasting is a powerful capability of numpy, that, within certain rules, allows operations (such as addition) on two arrays whose shapes do not match.
Numpy enables this by repeating values along dimensions of length 1, and appending such dimensions if necessary.
Broadcasting enables efficient and concise operations on arrays without Python loops and without explicitly inserting or repeating dimensions.
If you are not familiar with broadcasting, we strongly recommend to read through out the very well written numpy tutorial on it (https://numpy.org/devdocs/user/basics.broadcasting.html).

In the following, we will provide a short recap of broadcasting.

First, print things out again so we know what we are talking about 

```python
##### DO NOT CHANGE #####
print("Matrix A: (Shape {})".format(A.shape))
print(A)
print("Vector v: (Shape {})".format(v.shape))
print(v)
print("Vector w: (Shape {})".format(w.shape))
print(w)

##### DO NOT CHANGE #####
```

Let's try something weird and and subtract a vector from a matrix

```python
##### DO NOT CHANGE #####
print(A - v)

##### DO NOT CHANGE #####
```

As we see, numpy did not give us an error but did something - it subtracted the vector from each row of the matrix. This mechanism is called **broadcasting** and it will be your new best friend (you will sometimes also hate it)! 

In general, if the shapes of two numpy arrays do not match, but the non-matching dimension is "1" for either of the arrays the values along that dimension will be broadcasted. Here, numpy implicitly added a axis of size 1 to the vector "v" and broadcasted the values of "v" to all rows of A. This implicity broadcasting can be a bit counter intuitive and is a common source of errors.

Yet, broadcasting is a very important tool to efficiently combine data in the ways you need to and you will see many applications throughout the exercises. As a general tip: **never** use "for-loops" to manipulate your arrays but try to find a way using broadcasting - it is just much much faster. We might deduce points in the exercises for avoidable for-loops.

So let's do the broadcasting more explicitly: 

```python
##### DO NOT CHANGE #####
v1 = v[np.newaxis, :] # equivalent to v[None, :] or np.reshape(v, (1, 3))
print("v1:", v1)
print("Shape:", v1.shape)
print("Subtraction: \n", A - v1)

##### DO NOT CHANGE #####
```

If we want to subtract the vector from each column instead of each row, we need to broadcast along the other axis! 

```python
##### DO NOT CHANGE #####
v2 = v[:, np.newaxis] # equivalent to v[:, None] or np.reshape(v, (3, 1))
print("v2:\n", v2)
print("Shape:", v2.shape)
print("Subtraction: \n", A - v2)

##### DO NOT CHANGE #####
```

## Exercise 3: Gaussian Distributions
The following code samples 1000 samples from a Gaussian distribution with mean $\mu = -2$ and standard deviation $\sigma=0.5$. <br>

To visualize the samples, we use matplotlib, a python package for plotting.

```python
##### DO NOT CHANGE #####
import matplotlib.pyplot as plt

# mean and standard deviation
mu = -2
sigma = 0.5

# use a fixed seed to always get the same samples
np.random.seed(0)

# samples an array of shape (1000,) from the normal distribution with mean mu and standard deviation sigma
samples = np.random.normal(mu, sigma, size=1000)

# plot the histogram of the samples over 40 bins
count, bins, ignored = plt.hist(samples, 40, density=True)

# show the plot
plt.show()


##### DO NOT CHANGE #####
```

### Task 3.1: Density Function

To verify that the generated samples actually follow a normal distribution, we want to plot the density of the distribution on top of the samples.

Implement the function `compute_normal_density`, which takes in a numpy array of locations `x`, a (scalar) mean `mu` and a (scalar) standard deviation `sigma` and computes the density of the normal distribution
$$
f(x) = \frac{1}{\sqrt{2 \pi} \sigma} \exp\left(-\frac{(x - \mu)^2}{2\sigma^2}\right).
$$
Remember that all scalar numpy operations (like scalar division, exponentials, and special functions such as `np.sqrt`) operate element-wise when given an array. You do not need any for loops!

<!-- grade_id: task_3_1, points: 0, grade: false, solution: true -->
```python
def compute_normal_density(x: np.ndarray, mu: np.ndarray, sigma: np.ndarray) -> np.ndarray:
    density = np.zeros_like(x)

    # Compute the density function of the normal distribution for the locations x. Overwrite the density variable.
    # Do not use any (explicit) loops!
    # Use the numpy functions np.sqrt (square root) and np.exp (exponential function)
    # Numpy also provides the constant np.pi, use it!

    # YOUR CODE HERE
    raise NotImplementedError()

    return density
```

<!-- grade_id: test_task_3_1, points: 0, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_3_1 - possible points: 0

# plot the histogram of the samples 
count, bins, ignored = plt.hist(samples, 40, density=True)

# plot the density function
densities = compute_normal_density(bins, mu, sigma)
assert densities.shape == bins.shape
plt.plot(bins, densities, linewidth=2, color='r')

# show the plot
plt.show()


##### DO NOT CHANGE #####
```

## Exercise 4: Solving a Linear Equation

You are given the linear equation
$$
Ax=v.
$$

### Task 4.1
We want to find the vector $x$. Assume that the matrix $A$ and the vector $v$ are given as in Exercise 1.1.
Find the solution for $x$. 

Use the function `np.linalg.solve` from numpy - this should always be preferred over direct matrix inversion.

<!-- grade_id: task_4_1, points: 0, grade: false, solution: true -->
```python
x = None

# Solve for x here
# YOUR CODE HERE
raise NotImplementedError()

print("Solution: x =", x)

```

<!-- grade_id: test_task_4_1, points: 0, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_4_1 - possible points: 0

print(f"Test: A@x = {A@x} is the same as v = {v}")


##### DO NOT CHANGE #####
```

## Closing Remarks

NumPy provides easy and efficient ways to deal with vectors, matrices and tensors. Linear algebra functionality like
the function introduced above, together with more advanced techniques such as broadcasting
(https://numpy.org/doc/stable/user/basics.broadcasting.html ), are the key to efficient programming in Python.
If you want to do any practical Machine Learning you need to familiarize yourself with those techniques.
We will use and introduce more of this in the following exercises but also expect you to use it for your submission.
So use them whenever possible and, most importantly, avoid unnecessary for loops.

