# Homework 1: Regression & Classification
**Submission Deadline: May 22, 2026, Friday, 9:30**

---

> **Machine Learning - Foundations and Algorithms** in Summer Semester 2026
>
> - Prof. Gerhard Neumann
> - Instructors for this exercise: Florian Seligmann (florian.seligmann@kit.edu), Serge Thilges (serge.thilges@kit.edu), and Emiliyan Gospodinov (emiliyan.gospodinov@kit.edu)
> 
> ⚠️ For general questions about the exercises, please post in the Ilias forum so we can answer them once for all students. Only use email for individual questions whose answers are not relevant to all students.

---

## Homework Info

This exercise is about simple (probabilistic) methods for regression and classification.
For regression, you will implement ridge regression, and for classification a generative naive Bayes classifier and a discriminative logistic regression model.

This homework contains both theory tasks and coding tasks. **Only the coding and multiple-choice tasks will be graded, all other tasks are marked with "(Not Graded)".** We still recommend to complete the theory tasks as they will prepare you for the exam, and we will discuss them in the exercise session.

## Group Submission

- The exercise sheets can be submitted in **groups of up to 3 students**.
- **Each person in the group must upload the edited notebook to Ilias**. It is not sufficient if only one person from the group does this. If you forgot to upload but at least on other group member upholds his solution, you will receive 2/3 of the points the **first** time this happens, and zero points if it **happens again**.
- **The submission deadline is strict**. Late submissions receive zero points. **We do not accept submissions via e-mail**. 
- It is possible to join a new group during the semester if your own group dissolves early. Each group must upload their own solution, and we will check submissions for duplicates.
- You can use the team building forum (https://ilias.studium.kit.edu/ilias.php?baseClass=ilrepositorygui&cmdNode=xs:lz:5y&cmdClass=ilBookingObjectGUI&ref_id=2930523) to team up. If you already have group members or want to submit alone, you do not need to use the forum.


Groups are automatically recorded for each exercise, so **please enter the u-identifiers of your group members in the *next* cell.** If your group consists of only 2 students or if you are submitting alone, leave the remaining fields empty.

Here is an **example** for a group consisting of uabcd and uefgh:

_U-identifiers of group members:_

_Member 1: uabcd_

_Member 2: uefgh_

_Member 3:_

Enter your group members in the *next* cell!

U-identifiers of group members:

Member 1:

Member 2:

Member 3:

## Auto-grading

We use an auto-grading system that automatically analyzes your submitted Jupyter Notebooks and checks them for correctness using hidden tests. These tests determine the points you receive for the exercise sheet.

To ensure auto-grading works smoothly, please note the following:

- The notebook must have the filename "ex_01_regression_classification.ipynb"
- **Upload only the Jupyter Notebook** on Ilias as a single file (not as a zip!). **Do not include any additional files**.
- Before submitting a notebook, test that everything runs without errors from start to finish.
- **Cells marked with "##### DO NOT CHANGE #####" must not be edited or deleted**
  - There may be seemingly empty cells that are also marked with "##### DO NOT CHANGE #####". These must also not be edited or deleted.
  - If you do modify them, auto-grading will not work and you will receive no points.
  - We will be strict about this and make no exceptions if someone modifies cells clearly marked as read-only!
- Your solution must be entered in the correct cell (marked with "# YOUR CODE HERE").
  - Please **delete the NotImplementedError**
- Generally, **do not delete any cells** and **do not add any cells**. The cells where your solution should be entered already exist (marked with "# YOUR CODE HERE").
- The Jupyter notebooks have inline tests (visible to you) that check your result for basic correctness.
  - These are primarily for you to identify and correct errors.
  - However, the inline tests you can see in the notebook are not the tests used for grading!
  - Passing the inline tests is a necessary but not sufficient condition to receive points!

**If you fail to comply with these rules, your submission may not be graded and you may receive no points for this exercise!**

## Exercise 1: Linear Regression

### Task 1.1: (Multiple-Choice) Matrix Vector Calculus (1 Point)
We'll start with some vector calculus warmup. Let

$$
X, Y \in \mathbb{R}^{d \times d}
$$

and

$$
\boldsymbol{q}, \boldsymbol{z} \in \mathbb{R}^d.
$$

Given the following element-wise expression, choose the equivalent matrix expression:

$$
    g = \alpha \sum_i \sum_j \sum_k z_k \, x_{ij} \, q_i \, y_{jk}
$$

- a)  $\displaystyle g = \alpha \, \boldsymbol{q}^T X Y \boldsymbol{z}$

- b)  $\displaystyle g = \alpha \, \boldsymbol{q}^T X^T Y \boldsymbol{z}$

- c) $\displaystyle g = \alpha \, \boldsymbol{q}^T X Y^T \boldsymbol{z}$

- d)  $\displaystyle g = \alpha \, \boldsymbol{q}^T Y X \boldsymbol{z}$


In the next cell, write your answer as a Python string with a single lowercase letter (as only one answer is correct) in the `answer_1_1` variable and remove raise NotImplementedError(). 

<!-- grade_id: task_1_1, points: 0, grade: false, solution: true -->
```python
answer_1_1 = ""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_task_1_1, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_1_1 - possible points: 1

assert answer_1_1 in ("a", "b", "c", "d")


##### DO NOT CHANGE #####
```




### Task 1.2 (Theory): Ridge Regression Weight Derivation (Not Graded)
Derive the optimal solution of weights in Ridge Regression using matrix form, i.e., solve for the $\boldsymbol{w}$ that minizes the objective function
$$
    L = (\boldsymbol{y}-\boldsymbol{\Phi} \boldsymbol{w})^T(\boldsymbol{y}-\boldsymbol{\Phi} \boldsymbol{w}) + \lambda \boldsymbol{w}^T \boldsymbol{I} \boldsymbol{w}.
$$

Hint: You will need derivatives for vectors/matrices.

### Task 1.3 (Coding) Ridge Regression Weight Implementation (2 Points)

First, the data. Run the cell below to load and visualize both the training an test data that we'll use for this task.

```python
##### DO NOT CHANGE #####
import numpy as np
import matplotlib.pyplot as plt
from typing import Callable

# Load data

training_data = np.load('training_data.npy')
test_data = np.load('test_data.npy')

test_data_x = test_data[:, 0]
test_data_y = test_data[:, 1]

training_data_x = training_data[:, 0]
training_data_y = training_data[:, 1]

# Visualize data
plt.plot(test_data_x, test_data_y, 'or')
plt.plot(training_data_x, training_data_y, 'ob')
plt.xlabel('x')
plt.ylabel('y')
plt.legend(["Test Data", "Training Data"])

##### DO NOT CHANGE #####
```

As you see, the data is not not linear. We will use polynomial features of degreee 10 to model this nonlinearity. 

Looking at the above plot of the data, the features ($x$ values) span from -4 to 4.
For the polynomial terms of high degree, this will lead to very large values that introduce numerical problems.
We will therefore *standardize* the features: Normalize each feature to (empirical) zero mean and (empirical) unit variance.
In general terms, given a feature matrix $\boldsymbol{\Phi} \in \mathbb{R}^{N \times M}$, standardization transforms them according to
$$
    \boldsymbol{\tilde{\Phi}} = \frac{\boldsymbol{\Phi}(\boldsymbol{x}) - \boldsymbol{\mu}_{\Phi}}{\boldsymbol{\sigma}_{\Phi}}
$$
where $\mu_{\Phi, i} = \frac{1}{M} \sum_{j=1}^M \Phi_j$ is the empirical mean of the training features and $\sigma_{\Phi, i} = \sqrt{\frac{1}{M} \sum_{j=1}^N (\Phi_j - \mu_{\Phi, j})^2}$ is the empirical standard deviation of the training features.


The next cell defines utility functions for
- calculating polynomial features
- computing the mean and std of the features (training data) as normalizer
- normalize other data (test) features using the normalizer (mean and std)
- evaluating the model
- calculating the Mean Squarred Error for evaluation of the model

```python
##### DO NOT CHANGE #####
def get_polynomial_features(data: np.ndarray, degree: int) -> np.ndarray:
    """
    Function to create Feature Matrix. Extends the feature matrix according to
    the matrix form discussed in the lectures.

    :param data: data points you want to evaluate the polynomials,
                 shape: [n_samples] (we have 1-dim data)
    :param degree: degree of your polynomial, shape: scalar
    :return polynomial_features: shape [n_samples x (degree+1)]
    """
    polynomial_features = np.ones(data.shape)
    for i in range(degree):
        polynomial_features = np.column_stack((polynomial_features, data ** (i + 1)))
    return polynomial_features


def get_mean_std_features(polynomial_features: np.ndarray) -> tuple[np.ndarray, np.ndarray]:
    """
    Function for calculating the mean and standard deviation of the features
    :param polynomial_features: shape: [n_samples x (degree+1)]
    :return mean_feat: mean vector of the features,
            shape:[1 x (degrees+1)]
    :return std_feat: standard deviation (for each dimension in feature matrix),
                      shape: [1 x (degrees+1)] 
    """
    mean_feat = np.mean(polynomial_features, axis=0, keepdims=True)
    mean_feat[:, 0] = 0.0 # we don't want to normalize the bias
    std_feat = np.std(polynomial_features, axis=0, keepdims=True)
    std_feat[:, 0] = 1.0 # we don't want to normalize the bias
    return mean_feat, std_feat


def normalize_features(polynomial_features: np.ndarray,
                       mean_train_features: np.ndarray,
                       std_train_features: np.ndarray) -> np.ndarray:
    """
    Normalize features
    :param polynomial_features:  features to be normalized,
                 shape: [n_samples x (degree+1)]
    :param mean_train_features: mean of the feature matrix of the training set,
                 shape: [1 x (degrees+1)]
    :param std_train_features: std of the feature matrix of the training set,
                 shape: [1 x (degrees+1)]
    :return norm_feat: normalized features, shape: [n_samples x (degree+1)]
    """

    # note: features: (n_samples x n_dims),
    #       mean_train_features: (1 x n_dims),
    #       std_train_features:  (1 x n_dims)
    #       due to these dimensionalities we can do element-wise operations.
    #       By this we normalize each dimension independently
    norm_feat = (polynomial_features - mean_train_features) / std_train_features
    return norm_feat


def eval(Phi: np.ndarray, w: np.ndarray)-> np.ndarray:
    """
    Evaluate the models

    :param Phi: Feature matrix, shape: [n_samples x (degree+1)]
    :param w: weight vector, shape: [degree + 1]
    :return : predictions, shape [n_samples] (we have 1-dim data)
    Evaluates your model
    """
    return np.dot(Phi, w)


def mse(y_target: np.ndarray, y_pred: np.ndarray)-> np.ndarray:
    """
    :param y_target: the target outputs,
            shape: [n_samples] (here 1-dim data)
    :param y_pred: the predicted outputs,
            shape: [n_samples](we have 1-dim data)
    :return : The Mean Squared Error, shape: scalar
    """
    diff = y_target - y_pred
    return np.sum(diff ** 2, axis=0) / y_pred.shape[0]

##### DO NOT CHANGE #####
```

Your task is to implement the function `calc_weights_ridge` that computes the weights for ridge regression. Recall that the weights are given by
$$
\boldsymbol{w} = (\boldsymbol{\Phi} ^T \boldsymbol{\Phi} + \lambda \boldsymbol{I} )^{-1} \boldsymbol{\Phi}^T \boldsymbol{y}
$$
where $\boldsymbol{\Phi}$ is the feature matrix containing the (polynomial) normalized features.

<!-- grade_id: task_1_3, points: 0, grade: false, solution: true -->
```python
def calc_weights_ridge(Phi: np.ndarray, y:np.ndarray, ridge_factor:float) -> np.ndarray:
    """
    :param Phi: Feature Matrix, shape: [n_samples x (degree+1)]
    :param y: Output Values, [n_samples] (we have 1-dim data)
    :param ridge_factor: lambda value, shape: scalar
    :return : The weight vector, calculated according to the equation shown before,
            shape: [degrees +1]
    """
    weights = np.zeros((Phi.shape[1],))

    # compute the weights here
    # use np.linalg.solve for solving the linear equation and `np.eye` to create an identity matrix.
    # overwrite the `weights` variable
    # YOUR CODE HERE
    raise NotImplementedError()

    return weights

```

The next cell will compute the normalized training and test features and then use your function to find the optimal ridge regression weights for various values of $\lambda$. It will then plot the mean squared error of your model on the training and test data.

<!-- grade_id: test_task_1_3, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_1_3 - possible points: 2

# Compute polynomial features
train_features = get_polynomial_features(training_data_x, 10)
test_features = get_polynomial_features(test_data_x, 10)

# Get empirical mean and standard deviation of the training features
mean_train_feat, std_train_feat = get_mean_std_features(train_features)

# Standardize training and test features
norm_train_features = normalize_features(train_features, mean_train_feat, std_train_feat)
norm_test_features = normalize_features(test_features, mean_train_feat, std_train_feat)

# Compute optimal weights for various lambda
ridge_factors = [1e-6, 1e-3, 1, 3, 5, 10,20,30,40, 50, 1e2, 1e3, 1e5]
weights_ridge = []

for lambda_val in ridge_factors:
    weights = calc_weights_ridge(norm_train_features, training_data_y, lambda_val)
    assert weights.shape == (norm_train_features.shape[1],), f"Expected shape ({norm_train_features.shape[1]},) but got {weights.shape}"
    weights_ridge.append(weights)
print("✓ Shapes are correct")

# Get predictions of the model on training and test features for evaluation
y_training_ridge = []
y_test_ridge = []

for w in weights_ridge:
    y_training_ridge.append(eval(norm_train_features, w))
    y_test_ridge.append(eval(norm_test_features, w))

# Compute the MSE on training and test data
training_error_ridge = []
test_error_ridge = []

for i in range(len(y_training_ridge)):
    training_error_ridge.append(mse(training_data_y, y_training_ridge[i]))
    test_error_ridge.append(mse(test_data_y, y_test_ridge[i]))

# Plot the MSE against lambda
error_fig_ridge = plt.figure()
plt.figure(error_fig_ridge.number)
plt.title("Error Plot Ridge Regression")
plt.xlabel("$\\lambda$")
plt.ylabel("MSE")
x_axis = ["$1e-{6}$", "$1e-{3}$", "$1$", "$3$", "$5$","$10$","$20$","$30$","$40$","$50$",
          "$1e2$", "$1e3$", "$1e5$"]
plt.yscale('log')
plt.plot(x_axis, training_error_ridge, 'b')
plt.plot(x_axis, test_error_ridge, 'r')

# let's find the index with the minimum training error
min_error_idx = np.argmin(test_error_ridge)
plt.plot(x_axis[min_error_idx], test_error_ridge[min_error_idx], 'xg')
plt.legend(['Training Error', 'Test Error', 'Min Test Error'])
plt.show()


##### DO NOT CHANGE #####
```

The optimal $\lambda$ among the tested values seems to be $\lambda = 10$. Note that this is a very high value. In practice, we would typically expect very small values, often in the range of $10^{-6}$ or smaller. The reason why we need such a large $\lambda$ is the high degree of the polynomial features. A degree of 10 makes the model highly expressive and leads to massive overfitting without appropriate regularization. In practice, we would not choose a model where we know from the beginning that it is highly overfitting.

Finally, let us visualize the model's predictions for the optimal value of $\lambda$.

```python
##### DO NOT CHANGE #####
x = np.linspace(-5, 5, 100)
new_features = get_polynomial_features(x, 10)
new_norm_feat = normalize_features(new_features, mean_train_feat, std_train_feat)
y_pred = eval(new_norm_feat, weights_ridge[min_error_idx])

plt.plot()
plt.plot(test_data_x, test_data_y, 'or')
plt.plot(training_data_x, training_data_y, 'ob')
plt.plot(x, y_pred)
plt.legend(["test_data", "training_data", "inference"])

##### DO NOT CHANGE #####
```

### Task 1.4 (Theory): Error Plot (Not Graded)

In the above error plot, analyze the relationship between the training/test error and different values of $\lambda$, as well as the reasons behind the observed behavior. Note that the y-axis in the error plot is in log-space: Small changes in the y-axis mean large differences in the error.

## Exercise 2: Probability Basics and Naive Bayes Classifier

Now we will turn to a classification task. Let's start by loading a very simple toy dataset, the "two moons":

```python
##### DO NOT CHANGE #####
import numpy as np
import matplotlib.pyplot as plt

data = dict(np.load("two_moons.npz", allow_pickle=True))
samples = data["samples"]
labels = data["labels"]

c0_samples = samples[labels == 0]  # class 0: all samples with label 0
c1_samples = samples[labels == 1]  # class 1: all samples with labe 1

plt.figure("Data")
plt.scatter(x=c0_samples[:, 0], y=c0_samples[:, 1], label="c0")
plt.scatter(x=c1_samples[:, 0], y=c1_samples[:, 1], label="c1")
plt.legend()
plt.show()

##### DO NOT CHANGE #####
```

The next cell defines some utility plotting functions.

```python
##### DO NOT CHANGE #####
def draw_2d_gaussian(mu: np.ndarray, sigma: np.ndarray, plt_std: float = 2, *args, **kwargs) -> None:
    (largest_eigval, smallest_eigval), eigvec = np.linalg.eig(sigma)
    phi = -np.arctan2(eigvec[0, 1], eigvec[0, 0])

    plt.scatter(mu[0:1], mu[1:2], marker="x", *args, **kwargs)

    a = plt_std * np.sqrt(largest_eigval)
    b = plt_std * np.sqrt(smallest_eigval)

    ellipse_x_r = a * np.cos(np.linspace(0, 2 * np.pi, num=200))
    ellipse_y_r = b * np.sin(np.linspace(0, 2 * np.pi, num=200))

    R = np.array([[np.cos(phi), np.sin(phi)], [-np.sin(phi), np.cos(phi)]])
    r_ellipse = np.array([ellipse_x_r, ellipse_y_r]).T @ R
    plt.plot(mu[0] + r_ellipse[:, 0], mu[1] + r_ellipse[:, 1], *args, **kwargs)

# plot grid for contour plots
plt_range = np.arange(-1.5, 2.5, 0.01)
plt_grid = np.stack(np.meshgrid(plt_range, plt_range), axis=-1)
flat_plt_grid = np.reshape(plt_grid, [-1, 2])
plt_grid_shape = plt_grid.shape[:2]

##### DO NOT CHANGE #####
```

In the following, we model the class conditional distributions $p(x~|~c)$ as Gaussian, the class prior $p(c)$ as Bernoulli, and apply Bayes' rule to compute the class posterior p(c~|~x):

$$
p(c~|~x) = \frac{p(x~|~c)\,p(c)}{p(x)} \propto p(x~|~c)\,p(c)
$$
where
$$
p(x~|~c) = \mathcal{N}(x;~\boldsymbol{\mu}_c, \boldsymbol{\Sigma}_c).
$$
is a multinomial normal distribution.
Note that we can ignore $p(x)$ as it is constant in $c$.

### Task 2.1 (Theory): Derivation of Maximum Likelihood Estimator (Not Graded)

We will use a maximum likelihood estimator to find $\boldsymbol{\mu}_c$ and $\boldsymbol{\Sigma}_c$.

The maximum likelihood estimator for the multinomial normal distribution is given by
$$
\boldsymbol{\mu} = \dfrac{1}{N} \sum_{i}^N \boldsymbol{x}_i \quad \quad \boldsymbol{\Sigma} = \dfrac{1}{N} \sum_{i}^N (\boldsymbol{x}_i - \boldsymbol{\mu}) (\boldsymbol{x}_i - \boldsymbol{\mu})^T.
$$

Recall that the density of the multinomial normal distribution $\mathcal{N}(\boldsymbol{x};~\boldsymbol{\mu}_c, \boldsymbol{\Sigma}_c)$ is given by
$$
\mathcal{N}\left(\boldsymbol{x} | \boldsymbol{\mu}, \boldsymbol{\Sigma} \right) = \dfrac{1}{\sqrt{\det \left(2 \pi \boldsymbol{\Sigma}\right)}} \exp\left( - \dfrac{(\boldsymbol{x}-\boldsymbol{\mu})^T \boldsymbol{\Sigma}^{-1} (\boldsymbol{x}-\boldsymbol{\mu})}{2}\right).
$$

Your task is to derive the above maximum likelihood estimator.

This derivations involves some matrix calculus. Matrix calculus is a bit like programming, you google the stuff you need and then plug it together in the right order. Good resources for such rules are the "matrix cookbook" (https://www.math.uwaterloo.ca/~hwolkowi/matrixcookbook.pdf) and the Wikipdia article about matrix calculus (https://en.wikipedia.org/wiki/Matrix_calculus ). State all rules you use explicitly (except the ones given in the hints below).

**Remark** There are different conventions of how to define a gradient (as column-vector or row-vector). This results in different ways to write the Jacobian and thus different, usually transposed, matrix calculus rules:

- In the lecture we define the gradient as column-vector
- In the Wikipedia article this convention is referred to as "Denominator Layout". It also contains a nice explanation of the different conventions for the gourmets among you ;)
- The Matrix Cookbook uses the same convention (gradient as column vector)
- Please use the column vector convention here

**Hint** Here are two of those rules that might come in handy:
$$
\frac{\partial \log \det(\boldsymbol{X})}{\partial\boldsymbol{X}} = \boldsymbol{X}^{-1}
$$
and, for symmetric matrices $\boldsymbol{A}$ (remember that covariance matrices are always symmetric)
$$
\frac{\partial \boldsymbol{x}^T\boldsymbol{A} \boldsymbol{x}}{\partial \boldsymbol{x}} = 2\boldsymbol{A}\boldsymbol{x}
$$
There is one rule missing to solve the exercise. You need to find it yourself. (Hint hint: Look in the matrix cookbook, chapter 2.2 ;) )

The next cell implement the density and maximum likelihood estimator for the multivariate normal distribution in Python. 
Note that this is only for educuational puposes; in practice, you would never implement it like this, but stay in the log-domain (i.e., only work with log-probabilities).
Also for numerically stable implementations of the multivariate normal density the symmetry and positive definitness of the covariance should be exploited by working with it's Cholesky decomposition $\boldsymbol{\Sigma} = \boldsymbol{L}\boldsymbol{L}^T$, where $\boldsymbol{L}$ is a lower triangular matrix.

```python
##### DO NOT CHANGE #####
def mvn_pdf(x: np.ndarray, mu: np.ndarray, sigma: np.ndarray) -> np.ndarray:
    """
    Density of the Multivariate Normal Distribution
    :param x: samples, shape: [N x dimension]
    :param mu: mean, shape: [dimension]
    :param sigma: covariance, shape: [dimension x dimension]
    :return p(x) with p(x) = N(mu, sigma) , shape: [N]
    """
    norm_term = 1 / np.sqrt(np.linalg.det(2 * np.pi * sigma))
    diff = x - np.atleast_2d(mu) # Same to mu[None]
    exp_term = np.sum(np.linalg.solve(sigma, diff.T).T * diff, axis=-1)
    return norm_term * np.exp(-0.5 * exp_term)

def mvn_mle(x: np.ndarray) -> tuple[np.ndarray, np.ndarray]:
    """
    Maximum Likelihood Estimation of parameters for Multivariate Normal Distribution
    :param x: samples shape: [N x dimension]
    :return mean (shape: [dimension]) und covariance (shape: [dimension x dimension]) that maximize likelihood of data.
    """
    mean = 1 / x.shape[0] * np.sum(x, axis=0)
    diff = x - mean
    cov = 1 / x.shape[0] * diff.T @ diff
    return mean, cov

##### DO NOT CHANGE #####
```

### Task 2.2 (Coding): Generative Classifier (2 Points)

In the enxt cell, we compute the posterior distribution $p(c~|~x)$ for a vector of samples $x$.

Fill in the computation of the marginal class likelihood $p(c=1)$ and the posterior computation in `compute_posterior`. Use `mvn_pdf` to compute the density of the multivariate normal distribution.
For the posterior, you'll need to compute the conditional likelihoods $p(x~|~c)$ for $c=0$ and $c=1$ and the marginal data likelihood $p(x)$.

<!-- grade_id: task_2_2, points: 0, grade: false, solution: true -->
```python
# Fit Gaussian Distributions using the maximum likelihood estimator to samples from both classes
mu_c0, sigma_c0 = mvn_mle(c0_samples)
mu_c1, sigma_c1 = mvn_mle(c1_samples)

# Prior obtained by "counting" samples in each class
p_c0 = c0_samples.shape[0] / samples.shape[0]

# Compute the marginal class likelihood p(c=1) here
p_c1 = None
# YOUR CODE HERE
raise NotImplementedError()

def compute_posterior(
        samples: np.ndarray,
        p_c0: float, mu_c0: np.ndarray, sigma_c0: np.ndarray,
        p_c1: float, mu_c1: np.ndarray, sigma_c1: np.ndarray) \
        -> tuple[np.ndarray, np.ndarray]:
    """
    computes the posteroir distribution p(c|x) given samples x, the prior p(c) and the
    class conditional likelihood p(x|c)
    :param samples: samples x to classify, shape: [N x dimension]
    :param p_c0: prior probability of class 0, p(c=0)
    :param mu_c0: mean of class conditional likelihood of class 0, p(x|c=0) shape: [dimension]
    :param sigma_c0: covariance of class conditional likelihood of class 0, p(x|c=0) shape: [dimension x dimension]
    :param p_c1: prior probability of class 1 p(c=1)
    :param mu_c1: mean of class conditional likelihood of class 1 p(x|c=1) shape: [dimension]
    :param sigma_c1: covariance of class conditional likelihood of class 1, p(x|c=1) shape: [dimension x dimension]
    :return two arrays, p(c=0|x) and p(c=1|x), both shape [N]
    """
    p_x_given_c0 = None
    p_x_given_c1 = None
    p_x = None
    p_c0_given_x = None
    p_c1_given_x = None

    # YOUR CODE HERE
    raise NotImplementedError()

    return p_c0_given_x, p_c1_given_x
```

<!-- grade_id: test_task_2_2, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_2_2 - possible points: 2

# Let's do some simple tests and compute the classifier's accuracy

p_c0_given_x, p_c1_given_x = compute_posterior(samples, p_c0, mu_c0, sigma_c0, p_c1, mu_c1, sigma_c1)

assert p_c0_given_x.shape == (100,), p_c0_given_x.shape
assert p_c1_given_x.shape == (100,), p_c0_given_x.shape
print("✓ Shapes are correct")

# Prediction
predicted_labels = np.zeros(labels.shape)

# Set the decision boundary at 0.5 (this is arbitrary)
predicted_labels[p_c0_given_x >= 0.5] = 0.0  # is not strictly necessary since whole array already zero.
predicted_labels[p_c1_given_x > 0.5] = 1.0

# Evaluate
acc = (np.count_nonzero(predicted_labels == labels)) / labels.shape[0]
print("Accuracy:", acc)


##### DO NOT CHANGE #####
```

Lets look at the class likelihoods:

```python
##### DO NOT CHANGE #####
plt.title("Naive Bayes")
plt.scatter(x=samples[labels == 0, 0], y=samples[labels == 0, 1], c="blue")
draw_2d_gaussian(mu_c0, sigma_c0, c="blue")
plt.scatter(x=samples[labels == 1, 0], y=samples[labels == 1, 1], c="orange")
draw_2d_gaussian(mu_c1, sigma_c1, c="orange")

##### DO NOT CHANGE #####
```

And, in the next cell, the final posterior distribution for class 1, $p(c = 1 ~|~ x)$. The color indicates the posterior likelihood for the respective call and the black line indicates the decision boundary.

We achieve a train accuracy of 87%. For such a simple task that is clearly not great, but it nicely illustrates a problem with generative approaches: They usually depend on quite a lot of assumptions.

```python
##### DO NOT CHANGE #####
p_c0_given_x, p_c1_given_x = compute_posterior(flat_plt_grid, p_c0, mu_c0, sigma_c0, p_c1, mu_c1, sigma_c1)
p_c0_given_x = np.reshape(p_c0_given_x, plt_grid_shape)
p_c1_given_x = np.reshape(p_c1_given_x, plt_grid_shape)

plt.figure(figsize=(10, 4))
plt.subplot(1, 2, 1)
plt.contourf(plt_grid[..., 0], plt_grid[..., 1], p_c0_given_x, levels=10)
plt.colorbar()
# plot decision boundary
plt.contour(plt_grid[..., 0], plt_grid[..., 1], p_c0_given_x, levels=[0.0, 0.5], colors=["k", "k"])

plt.title("p($c_0$ | x)")
s0 = plt.scatter(c0_samples[..., 0], c0_samples[..., 1], color="blue")
s1 = plt.scatter(c1_samples[..., 0], c1_samples[..., 1], color="orange")
plt.legend([s0, s1], ["c0", "c1"])
plt.xlim(-1.5, 2.5)

plt.subplot(1, 2, 2)
plt.contourf(plt_grid[..., 0], plt_grid[..., 1], p_c1_given_x, levels=10)
plt.colorbar()
# plot decision boundary
plt.contour(plt_grid[..., 0], plt_grid[..., 1], p_c0_given_x, levels=[0.0, 0.5], colors=["k", "k"])
plt.title("p($c_1$ | x)")
s0 = plt.scatter(c0_samples[..., 0], c0_samples[..., 1], color="blue")
s1 = plt.scatter(c1_samples[..., 0], c1_samples[..., 1], color="orange")
plt.legend([s0, s1], ["c0", "c1"])

plt.xlim(-1.5, 2.5)
plt.show()


##### DO NOT CHANGE #####
```

### Task 2.3 (Multiple-Choice): Wrong Assumptions? (1 Point)

The generative classifier above achieves reasonable, but imperfect, accuracy. Which modeling assumption does not hold with respect to the given data distribution?

- a) The class prior $p(c)$ can be estimated from the relative class frequencies in the training data.
- b) Each class-conditional distribution $p(\boldsymbol{x} \mid c)$ is described by one multivariate Gaussian distribution.
- c) The covariance matrix may contain non-zero off-diagonal entries, allowing correlations between features.
- d) The posterior $p(c \mid \boldsymbol{x})$ can be computed from $p(\boldsymbol{x} \mid c)$ and $p(c)$ using Bayes' rule.

In the next cell, write your answer as a Python string with a single lowercase letter (as only one answer is correct) in the `answer_2_3` variable and remove raise NotImplementedError().


<!-- grade_id: task_2_3, points: 0, grade: false, solution: true -->
```python
answer_2_3 = ""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_task_2_3, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_2_3 - possible points: 1

assert answer_2_3 in ("a", "b", "c", "d")



##### DO NOT CHANGE #####
```

## Exercise 3: Logistic Regression with Batch Gradient Descent

In the recap sessions with Prof. Neumann we already saw (or will see) an implementation of a Discriminative Classifier using Logistic Regression. As you know, there is no closed for solution for the weights of a logistic regression model. Instead, we typically rely on iterative optimization using gradient descent. 

Here we are going to extend this to batch gradient descent.

We will use cubic polynomial features.
We start by implementing a few helper functions for affine mappings, the sigmoid function, and the negative Bernoulli log-likelihood, as well as the Bernoulli likelihood objective and its gradient.

```python
##### DO NOT CHANGE #####
def affine_features(x: np.ndarray) -> np.ndarray:
    """
    implements affine feature function
    :param x: inputs, shape: [N x sample_dim]
    :return inputs with additional bias dimension, shape: [N x feature_dim]
    """
    return np.concatenate([x, np.ones((x.shape[0], 1))], axis=-1)

def quad_features(x: np.ndarray) -> np.ndarray:
    """
    implements quadratic feature function
    :param x: inputs, shape: [N x sample_dim]
    :return squared features of x, shape: [N x feature_dim]
    """
    sq = np.stack([x[:, 0] ** 2, x[:, 1]**2, x[:, 0] * x[:, 1]], axis=-1)
    return np.concatenate([sq, affine_features(x)], axis=-1)

def cubic_features(x: np.ndarray) -> np.ndarray:
    """
    implements cubic feature function
    :param x: inputs, shape: [N x sample_dim]
    :return cubic features of x, shape: [N x feature_dim]
    """
    cubic = np.stack([x[:, 0]**3, x[:, 0]**2 * x[:, 1], x[:, 0] * x[:, 1]**2, x[:, 1]**3], axis=-1)
    return np.concatenate([cubic, quad_features(x)], axis=-1)

def quartic_features(x: np.ndarray) -> np.ndarray:
    """
    implements cubic feature function
    :param x: inputs, shape: [N x sample_dim]
    :return cubic features of x, shape: [N x feature_dim]
    """
    quartic = np.stack([x[:, 0]**4, x[:, 0]**3 * x[:, 1], x[:, 0]**2 * x[:, 1]**2, x[:, 0] * x[:, 1]**3, x[:, 1]**4], axis=-1)
    return np.concatenate([quartic, cubic_features(x)], axis=-1)

def quintic_features(x: np.ndarray) -> np.ndarray:
    """
    implements cubic feature function
    :param x: inputs, shape: [N x sample_dim]
    :return cubic features of x, shape: [N x feature_dim]
    """
    quintic = np.stack([x[:, 0]**5, x[:, 0]**4 * x[:, 1], x[:, 0]**3 * x[:, 1]**2, x[:, 0]**2 * x[:, 1]**3, x[:, 0] * x[:, 1]**4, x[:, 1]**5], axis=-1)
    return np.concatenate([quintic, cubic_features(x)], axis=-1)

def sigmoid(x: np.ndarray) -> np.ndarray:
    """
    the sigmoid function
    :param x: inputs
    :return sigma(x)
    """
    return 1 / (1 + np.exp(-x))

def bernoulli_nll(predictions: np.ndarray, labels: np.ndarray, epsilon: float = 1e-12) -> np.ndarray:
    """
    :param predictions: output of the classifier, shape: [N]
    :param labels: true labels of the samples, shape: [N]
    :param epsilon: small offset to avoid numerical instabilities (i.e log(0))
    :return negative log-likelihood of the labels given the predictions
    """
    return - (labels * np.log(predictions + epsilon) + (1 - labels) * np.log(1 - predictions + epsilon))

def objective_bern(weights: np.ndarray, features: np.ndarray, labels: np.ndarray) -> float:
    """
    bernoulli log-likelihood objective
    :param weights: current weights to evaluate, shape: [feature_dim]
    :param features: train samples, shape: [N x feature_dim]
    :param labels: class labels corresponding to train samples, shape: [N]
    :return average negative log-likelihood
    """
    predictions = sigmoid(features @ weights)
    return np.mean(bernoulli_nll(predictions, labels))

def d_objective_bern(weights: np.ndarray, features: np.ndarray, labels: np.ndarray) -> np.ndarray:
    """
    gradient of the bernoulli log-likelihood objective
    :param weights: current weights to evaluate, shape: [feature_dim]
    :param features: train samples, shape: [N x feature_dim]
    :param labels: class labels corresponding to train samples, shape [N]
    """
    res = np.expand_dims(sigmoid(features @ weights) - labels, -1)
    grad = features.T @ res / res.shape[0]
    return np.squeeze(grad)



##### DO NOT CHANGE #####
```

### Task 3.1 (Coding): Batch Gradient Descent (2 Points)

Implement the batch gradient descent update.

The code stores the weights for all iterations in the `weights` variable, so, in each iteration, you need to assign the updated weights to `weights[i + 1]`.

<!-- grade_id: task_3_1, points: 0, grade: false, solution: true -->
```python
def minimize_with_gradient_descent(features: np.ndarray, labels: np.ndarray, initial_weights: np.ndarray, schedule: Callable,
                      num_iterations: int):
    """
    :param features: all samples, shape: [N x feature_dim]
    :param labels: all labels, shape: [N]
    :param initial_weights: initial weights of the classifier, shape: [feature_dim * K]
    :param schedule: learning rate schedule (a callable function returning the learning rate, given the iteration
    :param num_iterations: number of times to loop over the whole dataset
    return "argmin", "min", logging info
    """


    # setup
    weights = np.zeros([num_iterations + 1, initial_weights.shape[0]])
    loss = np.zeros(num_iterations + 1)
    weights[0] = initial_weights
    loss[0]= objective_bern(weights[0], features, labels)

    for i in range(num_iterations):
        # get learning rate and perform batch gradient descent step
        lr = schedule(i)

        # Compute and write the updated weights
        weights[i + 1] = None
        # YOUR CODE HERE
        raise NotImplementedError()
        
        # log loss (on all samples, usually you should not use all samples to evaluate after each update step)
        loss[i + 1] = objective_bern(weights[i + 1], features, labels)

    return weights[-1], loss[-1], (weights, loss)

```

The next cell runs a few quick checks on your function.

<!-- grade_id: test_task_3_1, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: test_task_3_1 - possible points: 2

features = cubic_features(samples)
w_bce_test, l_test, l_info_test = minimize_with_gradient_descent(features, labels, np.zeros(features.shape[1]),
                                 schedule=(lambda t: 0.25),
                                 num_iterations=10)

assert w_bce_test.shape == (10,), w_bce_test.shape
assert l_test.shape == (), l_test.shape
assert len(l_info_test) == 2, len(l_info_test)
assert l_info_test[0].shape == (11, 10), l_info_test[0].shape
assert l_info_test[1].shape == (11,), l_info_test[1].shape
print("✓ Shapes are correct")


##### DO NOT CHANGE #####
```

Let's run the optimization loop and plot the resulting decision boundary.

You can freely edit the next cell. Feel free to play around with the various feature functions. As you'll see, the accuracy of this classifier is a lot better than that of the generative classifier!

```python
##### DO NOT CHANGE #####
# Generate Features from Data

# change this to play around with feature functions
#feature_fn = affine_features
# feature_fn = quad_features
feature_fn = cubic_features
# feature_fn = quartic_features
# feature_fn = quintic_features
features = feature_fn(samples)

num_iterations = 10000

w_bce, l, l_info = minimize_with_gradient_descent(features, labels, np.zeros(features.shape[1]),
                                 schedule=(lambda t: 0.25),
                                 num_iterations=num_iterations)
print("Final loss", l)

plt.figure()
plt.title("Cross Entropy Loss")
plt.grid("on")
plt.xlabel("Update Steps")
plt.ylabel("Negative Bernoulli Log-Likelihood")
plt.semilogy(l_info[1])

plt.figure()
plt.title("Bernoulli LL Solution")
pred_grid = np.reshape(sigmoid(feature_fn(flat_plt_grid) @ w_bce), plt_grid_shape)

plt.contourf(plt_grid[..., 0], plt_grid[..., 1], pred_grid, levels=10)
plt.colorbar()
#This is just a very hacky way to get a black line at the decision boundary:
plt.contour(plt_grid[..., 0], plt_grid[..., 1], pred_grid, levels=[0, 0.5], colors=["k"])

s0 = plt.scatter(c0_samples[..., 0], c0_samples[..., 1], color="blue")
s1 = plt.scatter(c1_samples[..., 0], c1_samples[..., 1], color="orange")
plt.legend([s0, s1], ["c0", "c1"])


##### DO NOT CHANGE #####
```